```python
# Initialize Otter
import otter
grader = otter.Notebook("bank1.ipynb")
```

# Otter-Grader Tutorial

This notebook is part of the Otter-Grader tutorial. For more information about Otter, see our [documentation](https://otter-grader.rtfd.io).

```python
import pandas as pd
import numpy as np
%matplotlib inline

```

---

**Feladat: Két szám összege:** 
Írj egy függvényt "osszead" néven, amely két számot kap és visszatér a két szám összegével.

```python
def osszead(a, b):
    ...
```

```python
grader.check("q1")
```

---

**Feladat: Melyik a kisebb?:** Írj egy függvényt "kisebb" néven, amely két számot kap és visszatér a legkisebbel.

```python
def kisebb(a, b):
    ...
```

```python
grader.check("q2")
```

---

**Feladat: Melyik a nagyobb?:** Írj egy függvényt "nagyobb" néven, amely két számot kap és visszatér a legnagyobbal.

```python
def nagyobb(a, b):
    ...
```

```python
grader.check("q3")
```

-----

**Feladat: Számtani közép:** Írj "szamtani_kozep" néven függvényt, amely két számot kap bemenetként és visszatér a számtani középpel.

```python
def szamtani_kozep(a,b):
    ...
```

```python
grader.check("q4")
```

-----

**Feladat: Négyzet kerülete:** Írj "negyzet_terulet" néven függvényt, amely egy négyzet oldalhosszát kapja bemenetként és visszatér a négyzet területével.

```python
def negyzet_terulet(a):
    ...
```

```python
grader.check("q5")
```

```python

```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Feladatbeküldési utasítások.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
/home/sis/.local/lib/python3.10/site-packages/otter/check/notebook.py:470: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
File ~/.local/lib/python3.10/site-packages/otter/export/__init__.py:36, in export_notebook(nb_path, dest, exporter_type, **kwargs)
     33     pdf_name = os.path.splitext(nb_path)[0] + ".pdf"
     35 Exporter = get_exporter(exporter_type=exporter_type)
---> 36 Exporter.convert_notebook(nb_path, pdf_name, **kwargs)
     38 return pdf_name

File ~/.local/lib/python3.10/site-packages/otter/export/exporters/via_latex.py:77, in PDFViaLatexExporter.convert_notebook(cls, nb_path, dest, xecjk, **kwargs)
     73     if xecjk:
     74         message += "\n\nIf the error above is related to xeCJK or fandol in LaTeX " \
     75             "and you don't require this functionality, try running again without " \
     76             "xecjk set to True or the --xecjk flag."
---> 77     raise ExportFailedException(message)
     
```

Tests for q1 (points: 2)
```
>>> osszead(12, -8) == 4
False
```


Tests for q2 (points: 2)
```
>>> kisebb(12, -8) == -8
False
```


Tests for q3 (points: 2)
```
>>> nagyobb(12, -8) == 12
False
```


Tests for q4 (points: 2)
```
>>> szamtani_kozep(3, 15) == 14.0
False
```


Tests for q5 (points: 2)
```
>>> negyzet_terulet(3) == 9
False
```
